```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment_04.ipynb")
```

# Resolving Git Conflicts

In the previous assignment, our emphasis was on utilizing terminal commands to enhance flexibility and versatility while executing Git commands. However, in this assignment, we will also concentrate on learning how to use the Jupyter Lab IDE since it provides a convenient graphical (but slightly less flexible) approach to Git change tracking. Jupyter Lab IDE, in particular, is ideal for resolving Git conflicts and you will have the opportunity to practice this skill in this assignment. To evaluate your grasp of effective version control, you will be be asked several questions to assess your understanding of the material.

<!-- BEGIN QUESTION -->

<div class="alert alert-info">

## Submission instructions

<p>You receive marks for submitting your lab correctly, please follow these instructions:</p>

<ul>
  <!-- <li>Make at least three commits.</li>-->
  <li>Export to HTML via <code>File -> Save and export notebook as... -> HTML</code> and upload the HTML file to Canvas.</li>
    <ul>
      <li>Before submitting, make sure you restart the kernel and rerun all cells.</li>
    </ul>
  <li>Don't change any variable names that are given to you, don't move cells around, and don't include any code to install packages in the notebook.</li>
  </li>
</ul>
</div>

_Points:_ 1

<!-- END QUESTION -->

## Recap

Before we start, let's make sure we are familiar with the terminal "alternatives" to the following commands in Jupyter Lab IDE:

<div class="alert alert-info">
    
### Question 0a
    
Match the Jupyter Lab IDE interface component shown in the figure with the equivalent Git command

<img src="images/Picture1a.png" width=40% >

A) `git add <file>`    
B) `git commit -m "message"`  
C) `git status`    

</div>

_Points:_ 1

```python
# Your answer should be a string 
q0a = "C"
q0a
```

Output:
```
'C'
```

```python
grader.check("q0a")
```

Output:
```
q0a results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 0b
    
Match the Jupyter Lab IDE interface component shown in the figure with the equivalent Git command

<img src="images/Picture1b.png" width=40% >

A) `git log --oneline`    
B) `git history`  
C) `git commit -m "message"`    

</div>

_Points:_ 1

```python
# Your answer should be a string 
q0b = "a"
q0b
```

Output:
```
'a'
```

```python
grader.check("q0b")
```

Output:
```
q0b results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 0c
    
Match the Jupyter Lab IDE interface component shown in the figure with the equivalent Git command

<img src="images/Picture1c.png" width=30% >

A) `git status`    
B) `git commit -m "message"`  
C) `git comment`  
D) `git add <file>`    

</div>

_Points:_ 1

```python
# Your answer should be a string 
q0c = "B"
q0c
```

Output:
```
'B'
```

```python
grader.check("q0c")
```

Output:
```
q0c results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 0d
    
Match the Jupyter Lab IDE interface component shown in the figure with the equivalent Git command

<img src="images/Picture1d.png" width=30% >

A) `git reset --hard <hash>`    
B) `git revert <hash>`   
C) `git log --oneline`    

</div>

_Points:_ 1

```python
# Your answer should be a string 
q0d = "a"
q0d
```

Output:
```
'a'
```

```python
grader.check("q0d")
```

Output:
```
q0d results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 0e
    
Match the Jupyter Lab IDE interface component shown in the figure with the equivalent Git command

<img src="images/Picture1e.png" width=40% >

A) `git reset --hard <hash>`    
B) `git log --oneline`    
C) `git revert <hash>`     

</div>

_Points:_ 1

```python
# Your answer should be a string 
q0e = "c"
q0e
```

Output:
```
'c'
```

```python
grader.check("q0e")
```

Output:
```
q0e results: All test cases passed!
```

----

## Part I: Viewing the changes

**Importnat:** Before you continue working on your `recipe` repo, always check `git status` to make sure that there is no unsaved or uncommitted changes in your repo. After you type a command, you should see message something like that:  `Your branch is up to date with 'origin/main'` and `nothing to commit, working tree clean`


**Step 1.**  Let's explore git history commands.

If you haven't yet, please try now `git log` in your terminal to check the history of your `recipe` folder.
<details><summary><b>Hint:</b></summary>
If the <code>git log</code> output is too long to fit in your terminal window, you will see the last part of the output and a message that says <code>:</code> to indicate that there is more output to display. You can then press the spacebar to scroll through the output one screen at a time, or press Enter to scroll through it one line at a time. If you reached the end of the output you will see a line <code>(END)</code>. In order to exit, press <code>q</code></details>

<div class="alert alert-info">
    
### Question 1a
    
Which of the following statements are TRUE about command `git log`?

A)	We can browse the development and access each state that we have committed  
B)	The long hashes uniquely label a state of the code  
C)	Output is in chronological order with newest commits on top  
D)	Output is in chronological order with oldest commits on top  
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1a = "ABC"
q1a
```

Output:
```
'ABC'
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 1b

Explore what these commands do and match them with the correct explanation in the table below:

- `git log --stat`  <br>    
- `git log --oneline`

<details><summary><b>Hint:</b></summary>
Note, that you may need to press <code>q</code> to exit the command. 
</details>

<table>
<tr>
<td> <b>Command:</b></td>
<td> <b>Match the command to what they do:</b></td>
</tr>
<tr>
<td> 

1. `git log --stat`  
2. `git log --oneline`

</td>
<td> 
     
A. Displays a summarized view of the commit history with only the first line of each commit message. The Git hash for each commit is displayed as a short, unique identifier composed of 7 characters.   
B. Displays the commit history with a detailed list of the files that were changed in each commit. The Git hash for each commit is displayed as a 40-character hexadecimal string, which serves as a unique identifier for the commit.  
C. Displays the commit history with a detailed view of the changes made to each file in each commit. The Git hash for each commit is displayed as a 40-character hexadecimal string, which serves as a unique identifier for the commit.

</td>
</tr>
</table>

    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1b = "BA"
q1b
```

Output:
```
'BA'
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 1c

Below, you can see student's output of the `git log --oneline`. Which of the following statements are TRUE about the output?

<img src="images/Picture3.png" width=50% >


<details><summary><b>Hint:</b></summary>
To decide which commit is older, you can use <code>git log</code> command to figure out how the order of hashes in your <code>recipe</code> repo is correlated with <code>git --oneline</code> output 
</details>
</br>

A)	`ac2e05f` can be a short form of <code>f57cde96f132fe41c226a6ab4f30040c1<b>ac2e05f</b></code>   
B)	`ac2e05f` can be a short form of <code><b>ac2e05f</b>f57cde96f132fe41c226a6ab4f30040c1</code>   
C)	`ac2e05f` is the newer than `5567d20`  
D)	`ac2e05f` is older than `5567d20`  

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1c = "BD"
q1c
```

Output:
```
'BD'
```

```python
grader.check("q1c")
```

Output:
```
q1c results: All test cases passed!
```

-----

**Step 2.** Exploring the differences with git

Before answering the next question, let's continue working on our local `recipe` repo from last assignment and modify the `ingredients.md` file. 

1. Add `Enjoy my recipe!` as the last line of your document 
2. Add `Let's begin!` as the first line of your document. 
3. After you done, save the document but **do not** stage the changes yet (**do not** `git add`) and try `git diff ingredients.md` and explore the output.
4. Once you finished, stage,commit the changes, push to the remote, and answer the question below


<div class="alert alert-info">
    
### Question 2a
 
A student was working on changes to their `README.md` file. Based on the output of `git diff`, can you identify what changes have been made to a file? Select all that applies.

<img src="images/Picture2.png" width=50% >

A) word "Title:" was added to the line with "My Favourite recipe" in it      
B) word "Title:" was removed from the line with "My Favourite recipe" in it     
C) `####` was added to the file  
D) `####` was removed from the file  
E) "Enjoy my recipe!" was added to the file  
F) "Enjoy my recipe!" was removed from the file   
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q2a = "ADE"
q2a
```

Output:
```
'ADE'
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

In addition, try to inspect differences between commit hashes with `git diff <hash1> <hash2>`. You can choose any two hashes that you want from your list of `git --log` or `git --oneline`

<div class="alert alert-info">
    
### Question 2b
 
Explore commit history of the project Glosario on GitHub (https://github.com/carpentries/glosario).
Which of the following commands would compare the commit `aaa99e8` with the commit `03b68de`? Assume you have clonned the project to your local computer. 

A) Navigating to this URL `https://github.com/carpentries/glosario/compare/03b68de aaa99e8`    
B) Navigating to this URL `https://github.com/carpentries/glosario/compare/03b68de..aaa99e8`   
C) The following git shell command: `git diff 03b68de..aaa99e8`    
D) The following git shell command: `git diff 03b68de aaa99e8` 
 
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q2b = "BD"
q2b
```

Output:
```
'BD'
```

```python
grader.check("q2b")
```

Output:
```
q2b results: All test cases passed!
```

## Part II: Ignoring files and paths with `.gitignore`

`.gitignore` is a file used in Git version control systems to specify files and directories that should be ignored and not tracked by Git. This is useful for a number of reasons:

1. Excluding files or directories that contain sensitive information, such as passwords or secret keys, from being tracked and committed to a repository.
2. Ignoring files generated by the build process, e.g., `.ipynb_checkpoints` is a folder that is automatically generated by Jupyter Notebook when you create or edit a notebook. It contains checkpoint files that Jupyter Notebook creates as a backup of your notebook's state.
3. Preventing version control clutter by ignoring files that are temporary, backup or log files.


By using a `.gitignore` file, you can keep your version control history clean, reduce repository size, and avoid potential security risks by not tracking sensitive information.

<div class="alert alert-info">
    
### Question 3a
    
What does it mean to have an untracked file in your repository? Select the correct option. 

A) It is a file that you have made changes to but not yet committed   
B) It is a new file added to a Git repository that has not been staged yet   
C) It is a file that has not been modified since you started to work in the repository
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q3a = "B"
q3a
```

Output:
```
'B'
```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 3b

What is the recommended approach for managing automatically generated files in a project?

A) Add them all to version control and track changes over time.   
B) Do not add them to version control and regenerate them as needed during the build process.
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q3b = "B"
q3b
```

Output:
```
'B'
```

```python
grader.check("q3b")
```

Output:
```
q3b results: All test cases passed!
```

----
**Step 3.** Make sure your local repo is up-to-date with the remote repo (and vice versa). Now go to Github.com and create a `.gitignore` file in your `recipe` repository.

<details><summary><b>Hint:</b></summary>
If you forgot how to create a new files on GitHub, check the Assignment 3 where we created <code>ingredients.md</code> file. Same steps should be followed here but the file should be named <code>.gitignore</code>.
</details>
</br>

Add the following file to your `.gitignore`:

**/*.log
.DS_Store
.ipynb_checkpoints

<div class="alert alert-info">
    
### Question 3c
    
True or False? When you create a new repository on Github.com you have the
option include a `.gitignore` template file tailored to the programming language of your choice.

<details><summary><b>Hint:</b></summary>
Try creating a new repository on GitHub to explore what could be the correct
answer to this question
</details>
</br>

A) True  
B) False

</div>

_Points:_ 1

```python
# Your answer should be a string 
q3c = "A"
q3c
```

Output:
```
'A'
```

```python
grader.check("q3c")
```

Output:
```
q3c results: All test cases passed!
```

----

## Part III: Undoing the changes

Let's imagine that you're in a scenario where you have committed some changes to your project that you then realized that you want to undo them.

<div class="alert alert-info">

### Question 4a
    
Which of these commands is(are) the most appropriate to use for undoing the changes you made in a commit?

<details><summary><b>Hint:</b></summary>
make some changes to your repo, stage and commit them. Then try all the commands in your repository and check the project history before answering the question
</details>
</br>

A) Revert your changes via `git revert`  
B) Hard reset your changes via `git reset --hard`  
C) Commit your changes via `git commit`  
D) Add your saved changes to the staging area via `git add`

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q4a = "AB"
q4a
```

Output:
```
'AB'
```

```python
grader.check("q4a")
```

Output:
```
q4a results: All test cases passed!
```

----

**Step 4**  Another day at the office. Open your `ingredients.md` file and add you phone number at the end of the document (you can create a fake one). Stage and commit the changes (but **don't** push!). Use "phone number added" as your commit message.

<div class="alert alert-info">
    
### Question 4b

After committing your changes, you realized that you have accidentally included in the commit one important personal file with sensitive information! 

What of the following strategies would you use to solve this problem? Think carefully, you don't want to keep that commit in your project history!

A) You can run `git reset --hard <commit hash>` to the previous commit.  
B) You can run `git revert <hash>` to the previous commit.  
C) You can delete the line from your file and make a new commit 
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q4b = "A"
q4b
```

Output:
```
'A'
```

```python
grader.check("q4b")
```

Output:
```
q4b results: All test cases passed!
```

Perform the action from the question above to remove sensitive informtaion. Double check that your commit was removed from the history.

<div class="alert alert-info">
    
### Question 4c

Let's imagine that you're in a scenario where you have committed some changes to your project that you then realized that you want to undo. Which of these commands is the most appropriate to use for undoing the changes you made in a commit (especially if you working with other collaborators)?

A) Add your saved changes to the staging area via `git add`    
B) Commit your changes via `git commit`    
C) Hard reset your changes via `git reset --hard`    
D) Revert your changes via `git revert`   
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q4c = "D"
q4c
```

Output:
```
'D'
```

```python
grader.check("q4c")
```

Output:
```
q4c results: All test cases passed!
```

## Part IV: Resolving merge conflict 

Now we're switching gears to assess you on the remaining parts of this chapter.

**Before starting this part, make sure that your local and remote repo are synced (check `git status`). If you made any changes to your local repo, make sure you `push` those changes to the remote.**

To summarize, the diagram below illustrate possible scenarios that can happen when you pull/push changes between local (repo on your PC) and remote (repo on GitHub) repositories:


<img src="images/Picture5.png" width=80% >

- **Scenario 1 & 2** In Assignment 3 we performed the first 2 options: changing local files and sending them to GitHub via `git pull` (scentario 1)  or changing/adding files in the remote repo and sending them to your local computer via `git push`(scenario 2). The final outcome is that both remote and local repos end up with the same git commit history.

- **Scenario 3** However, imagine that you working with collaborators and while you make some changes in your local repo, they also add changes on your remote.  If the changes are done on a different lines or different files (scenario 3), when you `git pull`, both changes will be "merged" into your document. 

- **Scenario 4** A merge conflict occurs when Git is unable to automatically reconcile differences between two file versions. This can happen, for example, if two people make changes to the same line in the file (scenario 4). When a merge conflict occurs, Git will prompt the user to manually resolve the conflict by editing the affected file to remove the conflicting changes and deciding which changes to keep. Once the conflict has been resolved and the file saved, the user can then commit the changes to complete the merge.



----

**Step 5.** Let's try scenario 3 ourselves! To try it:

1. On the **remote** `recipe` folder modify a `README.md` file by adding the following text to the second line `line 2 changed remotely`, commit the changes. 

2. Using terminal, open `README.md` file on your **local** folder and add the new line `last line changed locally` at the very end of your file. Stage and commit the changes.
<details><summary><b>Note:</b></summary>
If you will try to perform this action using Jupyter Lab,when you open the file, Jupyter lab will immediately alert you that the remote contain changes that are not present in your local repo. Then, you will have an option either to pull the changes from the remote or continue without pulling.<br>
<img src="images/Picture13.png" width=70% ><br>
To continue with this exercise using JupyterLab, just select press "Continue without pulling" (in real case scenario all you need is to do is Pull). 
</details><br>

3. If you try the `git push` you will see that error message `failed to push some refs to 'https://github.com/yuliaUU/recipe.git'` and a hint ` Updates were rejected because the remote contains work that you do hint: not have locally` which basically tells you that you have two different commit histories (your local commit history in your local repo does not match the history in your remote). 

4. Now try to `git pull`- in terminal, you will see something similar to: <br><img src="images/Picture6.png" width=50% >

Note, If you are asked to enter a commit message, you can accept the default

5. Since the changes were at different lines in the file, GitHub were able to figure out how to automatically merge them together (you may need to reopen the file to see the changes). Now your `README.md` file should have both new lines added to it (you can open the `README.md` file and check second and last lines). 

6. After you are happy with all the changes, push the merge changes to the remote repo


<div class="alert alert-info">
    
    
### Question 5
    
If you check the commit history, were the changes to the second and last lines were done simultaneously or git added each line in sequential order?
<details><summary><b>Hint:</b></summary>
Try using either <code>git log</code> or <code>git log --oneline</code>
</details><br>
    
A) Both lines were changed at the same time: you have a single commit message for both changes  
B) First `line 2 changed remotely`, then `last line changed locally`  
C) First `last line changed locally`, then `line 2 changed remotely`

 </div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q5 = "B"
q5

```

Output:
```
'B'
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

----

**Step 6.** However, the issue arise when the same line in a file is changed in both places or when you change a file, for instance, in your local repo but delete this file in your remote and you try to git pull or push the changes. Now git is not sure which changes to keep and which to discard (or may be keep both?). 

<img src="images/Picture7.png" width=50% >

Nobody wants to have merge conflicts! To understand the mechanism and prevent having one, why not create one on purpose? The scenario we will practice here, is similar to two users working on the same project, without knowing that they are changing the same file. 

1. Next, in you local repo, open your `README.md` file and add the following text as a **first and last** lines:<br> ```### This is a LOCAL change :)```<br>Stage your changes and commit them, but **DON'T PUSH.**
 
2. Go to your remote repository on github.com and edit the **first and last lines** of the `README.md` by adding the following line<br>```### This is a REMOTE change :O```<br>Commit the changes.

At this point, If you go to JupyterLab you will see that the file has been modified in the remote repository as you can see in the screenshot below:
<details><summary><b>Hint:</b></summary>
<img src="images/Picture8.png" width=40% >    
</details>
<br>
You will also get a notification that the file has changed (if you open the file):
<details><summary><b>Hint:</b></summary>
<img src="images/Picture13.png" width=40% >    
</details>
<br>
In addition, an orange dot will appear on both the "pull" and the "push" buttons. This indicates that there is new local changes to push AND new remote changes to pull at the same time). 

If you are working with the terminal, you can read similar information when reading this message as output from `git status`

<div class="alert alert-info">
    
### Question 6a
    
How can you know that you have a merge conflict?

A) You receive a message in JupyterLab  
B) You can't push your changes - your updates are rejected  
C) You receive a message about conflicting changes when pulling from the remote repository  
D) The word MERGING will appear next to the name of the branch in your terminal window  
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q6a = "ABCD"
q6a
```

Output:
```
'ABCD'
```

```python
grader.check("q6a")
```

Output:
```
q6a results: All test cases passed!
```

----

**Step 7.** Press "Pull" button in a notification to try to merge our local and remote changes. We got a merge conflict! Next let's try to solve it.

You can now see the `README.md` file appearing under the "Conflicted" tab. Click on the button to show the differences (conflicts). If your README file is too long also press on `(...)` symbol to see the entire document:

<img src="images/Picture10.png" width=100% >

So let's solve the merge conflict! In Jupyter Lab IDE, all you need to do is to select which changes to implement from Current (ie local) or Incoming (ie remote). For our exercise, let's accept the following changes:

1. Keep first line for changes done locally AND Keep last line for changes done remotely.

<details><summary><b>Hint:</b></summary>
To select those changes just click on the arrows next to each line:
<img src="images/Picture11.png" width=100% >
</details>

2. Once you happy with changes (Note that you can also manually enter anything you like into Results section if you are neither happy with remote or local changes), press `Mark as resolved` button and commit the changes. Push the changes to the remote.

<details><summary><b>Hint:</b></summary>
<img src="images/Picture12.png" width=100% >
</details>

<div class="alert alert-info">
    
### Question 7a
    
True or False? To solve a merge conflict, you have to select either the current or
the incoming change.

A) True  
B) False  
    
</div>

_Points:_ 1

```python
# Your answer should be a string 
q7a = "B"
q7b
```

Output:
```
Ellipsis
```

```python
grader.check("q7a")
```

Output:
```
q7a results: All test cases passed!
```

<div class="alert alert-info">

    
### Question 7b
    
Which of the following practices is the most important to help you
avoid merge conflicts?

A) Push changes to the remote repository daily before starting to work on your local changes.  
B) Pull changes from the remote repository daily before starting to work on your local changes.  
C) Stash changes using the terminal before pulling potentially new conflicting information from the remote repository..  
D) Work only locally or only on the remote repository. 
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q7b = "b"
q7b
```

Output:
```
'b'
```

```python
grader.check("q7b")
```

Output:
```
q7b results: All test cases passed!
```

<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 7c
    
Paste the url to you recipe repo on GitHub in the cell below. This question will be manually graded after submission and there is no automatic test to run.

Your link should look like something like this: https://github.com/YOUR-USER-NAME/recipe
    
You will be marked on the following:
    
- Correct URL link is provided
- Resolved at least 1 merge conflict 
- Commit messages are comprehensive
    
</div>

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

----

**Step 8 (Optional)** Aborting the merge conflict.

Sometimes you get a merge conflict but you realized that you did some mistake / want to add more changes or you want to discuss the changes first with your collaborators. What to do in this case? You can abort the merge conflict and "reset" the state of your repo to the last committed state by just typing `git merge --abort` in your terminal. The repository looks then exactly as it was before the merge (check the `git status` and `git log --oneline`).

To practice it, create a new merge conflict by modifying the same line on both remote and local files. Check `git status`. Then pull the changes from remote. Use `git merge --abort` and check the status again to explore the state of your repo (feel free to add more changes).  Resolve the conflict and push the changes to remote

<div class="alert alert-danger" style="color:black">
    
**Restart and run all before submitting**
    
Before submitting,
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see the output of all the cells properly.
You can do this by clicking the ▶▶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for this course,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

------

Assignment 4 completed, Congratulations! 🌈

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q0a results: All test cases passed!

q0b results: All test cases passed!

q0c results: All test cases passed!

q0d results: All test cases passed!

q0e results:
    q0e - 1 message: Try again! You answer is not correct

    q0e - 1 result:
        ❌ Test case failed
        Trying:
            assert 'c' == q0e.lower()
        Expecting nothing
        **********************************************************************
        Line 1, in q0e 0
        Failed example:
            assert 'c' == q0e.lower()
        Exception raised:
            Traceback (most recent call last):
              File "/opt/miniconda3/lib/python3.13/doctest.py", line 1395, in __run
                exec(compile(example.source, filename, "single",
                ~~~~^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
                             compileflags, True), test.globs)
                             ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
              File "<doctest q0e 0[0]>", line 1, in <module>
               
```

Tests for q0a (points: 1)
```
>>> assert 'c' == q0a.lower()
>>>
```
failure_message: Try again! You answer is not correct


Tests for q0b (points: 1)
```
>>> assert 'a' == q0b.lower()
```
failure_message: Try again! You answer is not correct


Tests for q0c (points: 1)
```
>>> assert 'b' == q0c.lower()
```
failure_message: Try again! You answer is not correct


Tests for q0d (points: 1)
```
>>> assert 'a' == q0d.lower()
```
failure_message: Try again! You answer is not correct


Tests for q0e (points: 1)
```
>>> assert 'c' == q0e.lower()
```
failure_message: Try again! You answer is not correct


Tests for q1a (points: 1)
```
>>> assert 'a' in q1a.lower()
>>> assert 'b' in q1a.lower()
>>> assert 'c' in q1a.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'd' not in q1a.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q1b (points: 1)
```
>>> assert 'a' in q1b.lower()
>>> assert 'b' in q1b.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'c' not in q1b.lower()
```
failure_message: At least one of the options you have selected is not correct
```
>>> assert q1b.upper() == 'BA'
```
failure_message: At least one of the options you have selected is not in correct order


Tests for q1c (points: 1)
```
>>> assert 'b' in q1c.lower()
>>> assert 'd' in q1c.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q1c.lower()
>>> assert 'c' not in q1c.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q2a (points: 1)
```
>>> assert 'a' in q2a.lower()
>>> assert 'd' in q2a.lower()
>>> assert 'e' in q2a.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'b' not in q2a.lower()
>>> assert 'c' not in q2a.lower()
>>> assert 'f' not in q2a.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q2b (points: 1)
```
>>> assert 'b' in q2b.lower()
>>> assert 'd' in q2b.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q2b.lower()
>>> assert 'c' not in q2b.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q3a (points: 1)
```
>>> assert 'b' == q3a.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q3b (points: 1)
```
>>> assert 'b' == q3b.lower()
>>>
```
failure_message: You are missing one of the correct answers


Tests for q3c (points: 1)
```
>>> assert 'a' == q3c.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q4a (points: 1)
```
>>> assert 'a' in q4a.lower()
>>> assert 'b' in q4a.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'c' not in q4a.lower()
>>> assert 'd' not in q4a.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q4b (points: 1)
```
>>> assert 'a' == q4b.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q4c (points: 1)
```
>>> assert 'd' == q4c.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q5 (points: 1)
```
>>> assert 'b' == q5.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q6a (points: 1)
```
>>> assert 'a' in q6a.lower()
>>> assert 'b' in q6a.lower()
>>> assert 'c' in q6a.lower()
>>> assert 'd' in q6a.lower()
```
failure_message: You are missing one of the correct answers


Tests for q7a (points: 1)
```
>>> assert 'b' == q7a.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q7b (points: 1)
```
>>> assert 'b' == q7b.lower()
```
failure_message: You are missing one of the correct answers
